---
title: 'Reshaping data'
description: ""
---

## Wide format data

```yaml
type: NormalExercise 
xp: 100 
key: 9dec5b91fb   
```


People often find it easier to record their data in wide format. In this format, each different data variable is presented in a separate column.


`@instructions`


`@hint`


`@pre_exercise_code`

```{r}
library('dplyr')
library('reshape2')
cars<-mtcars
```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


