---
title: 'Reshaping data'
description: ""
---

## Wide format data

```yaml
type: NormalExercise 
xp: 100 
key: 9dec5b91fb   
```


People often find it easier to record their data in wide format.


`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


