---
title: 'R Basics'
description: 'This is a template chapter.'
---

## An exercise title written in sentence case

```yaml
type: NormalExercise 
lang: r
xp: 100 
skills: 1
key: b6f33b15a3   
```


R is a command line driven program. It means that you enter your commands at a prompt (R Console), and an output is produced after you press "Enter".

Normally, people write their codes in a "Script.R" file, which is equivalent to a text file. That allows the user to save their code to work on it later.  

In the right hand side, you can see both the _Sript.R_ file, and the _R Console_. You can write your code in the _Srcipt.R_ window and see the output of your code in the _R Console_ by pressing the button **Run Code.**


`@instructions`
Please write the following command in the Script.R window and see the output in your R Console:

`` 2+2 ``

Once you see the output press **"Submit Answer"**

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`

```{r}
# Write your code below
___ + ___
```


`@sample_code`

```{r}

```


`@solution`

```{r}
2+2
```


`@sct`

```{r}

```


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 417dd4352b   
```


test


`@instructions`


`@hint`


`@pre_exercise_code`

```{r}
library('dplyr')
library('reshape2')
cars<-mtcars
```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 02d26e3683   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


